<!DOCTYPE html>
<!--
Copyright (c) 2012 Intel Corporation.

Redistribution and use in source and binary forms, with or without modification, 
are permitted provided that the following conditions are met:

* Redistributions of works must retain the original copyright notice, this list 
  of conditions and the following disclaimer.
* Redistributions in binary form must reproduce the original copyright notice, 
  this list of conditions and the following disclaimer in the documentation 
  and/or other materials provided with the distribution.
* Neither the name of Intel Corporation nor the names of its contributors 
  may be used to endorse or promote products derived from this work without 
  specific prior written permission.

THIS SOFTWARE IS PROVIDED BY INTEL CORPORATION "AS IS" 
AND ANY EXPRESS OR IMPLIED WARRANTIES, INCLUDING, BUT NOT LIMITED TO, THE 
IMPLIED WARRANTIES OF MERCHANTABILITY AND FITNESS FOR A PARTICULAR PURPOSE 
ARE DISCLAIMED. IN NO EVENT SHALL INTEL CORPORATION BE LIABLE FOR ANY DIRECT, 
INDIRECT, INCIDENTAL, SPECIAL, EXEMPLARY, OR CONSEQUENTIAL DAMAGES (INCLUDING, 
BUT NOT LIMITED TO, PROCUREMENT OF SUBSTITUTE GOODS OR SERVICES; LOSS OF USE, 
DATA, OR PROFITS; OR BUSINESS INTERRUPTION) HOWEVER CAUSED AND ON ANY THEORY 
OF LIABILITY, WHETHER IN CONTRACT, STRICT LIABILITY, OR TORT (INCLUDING 
NEGLIGENCE OR OTHERWISE) ARISING IN ANY WAY OUT OF THE USE OF THIS SOFTWARE, 
EVEN IF ADVISED OF THE POSSIBILITY OF SUCH DAMAGE.  
  
Authors:
        Zhang, Zhiqiang <zhiqiang.zhang@intel.com>

-->

<html>
  <head>
    <title>Tizen Test: bluetoothmanager_getdefaultadapter_base</title>
    <link rel="author" title="Intel" href="http://www.intel.com" />
    <link rel="help" href="https://developer.tizen.org/help/index.jsp?topic=%2Forg.tizen.help.web.api.device%2Ftizen%2Fbluetooth.html" />
    <meta name="flags" content="" />
    <meta name="assert" content="Check if getDefaultAdapter() return value is of type object" />
    <script src="../resources/testharness.js"></script>
    <script src="../resources/testharnessreport.js"></script>
    <script src="../resources/blacklist.js"></script>
  </head>
  <body>
    <div id="log"></div>
    <script type="text/javascript">
        function runtest()
        {
            //The api is in blacklist(not supported list)
            if (! is_platform_supported("BluetoothManager.getDefaultAdapter")) {
                try {
                    var adapter = tizen.bluetooth.getDefaultAdapter();
                    assert_true(typeof adapter == 'object', "getDefaultAdapter() return value is of type object");
                } catch (err) {
                    if (err.name == "NOT_SUPPORTED_ERROR") {
                        assert_true(true, "Mark the case as \"PASS\" because the platform actually doesn't support the method");
                    } else {
                        assert_false(true, "getDefaultAdapter() threw an exception: " + err.name + err.message);
                    }
                }

                return;
            }

            //The API is not in blacklist, then do normal tests
            try {
                var adapter = tizen.bluetooth.getDefaultAdapter();
                assert_true(typeof adapter == 'object', "getDefaultAdapter() return value is of type object");
            } catch (err) {
                assert_false(true, "getDefaultAdapter() threw an exception: " + err.name + err.message);
            }
        }

        test(runtest, document.title);
    </script>
  </body>
</html>
